# NYC Taxi Dataset

Viajes en taxi en la ciudad de Nueva York (sólo taxis amarillos) en los años 2015 y 2016.

**Taxis amarillos**: estos son los famosos taxis amarillos de la ciudad de Nueva York que prestan servicio exclusivamente mediante señas en la calle. El número de taxis está limitado por la cantidad finita de licencias emitidas por la TLC. Se accede a este modo de transporte colocándose en la calle y llamando con la mano a un taxi disponible. Las recogidas no están prearregladas.

**Vehículos de alquiler (FHV)**: el servicio FHV se accede mediante una reserva previa con un despachador o una compañía de limusinas. Estos FHV no están autorizados a recoger pasajeros mediante señas en la calle, ya que esos viajes no se consideran pre-reservados.

**Taxis verdes**: Street Hail Livery (SHL). El programa SHL permitirá a los propietarios de vehículos de alquiler (VTC) licenciar y equipar sus vehículos con la identidad de taxi verde del distrito, instalar taxímetros y máquinas para tarjetas de crédito y, en última instancia, obtener el derecho a aceptar señas en la calle además de viajes pre-reservados.

En este dataset se consideran únicamente los datos de taxis amarillos de los meses de enero de 2015 y de enero a marzo de 2016.

| Campo                   | Descripción                                                                                                                               |
|-------------------------|-------------------------------------------------------------------------------------------------------------------------------------------|
| `VendorID`              | Código que identifica al proveedor TPEP que suministró el registro (1 = Creative Mobile Technologies; 2 = VeriFone Inc.)                  |
| `tpep_pickup_datetime`  | Fecha y hora en que se activó el taxímetro.                                                                                               |
| `tpep_dropoff_datetime` | Fecha y hora en que se desactivó el taxímetro.                                                                                            |
| `passenger_count`       | Número de pasajeros en el vehículo (valor ingresado por el conductor).                                                                    |
| `trip_distance`         | Distancia recorrida en millas reportada por el taxímetro.                                                                                 |
| `pickup_longitude`      | Longitud del lugar donde se activó el taxímetro.                                                                                          |
| `pickup_latitude`       | Latitud del lugar donde se activó el taxímetro.                                                                                           |
| `RateCodeID`            | Código de tarifa vigente al final del viaje (1=Standard rate, 2=JFK, 3=Newark, 4=Nassau or Westchester, 5=Negotiated fare, 6=Group ride). |
| `store_and_fwd_flag`    | Indica si el registro se almacenó en memoria del vehículo antes de enviarlo (Y = store and forward; N = not a store and forward trip).    |
| `dropoff_longitude`     | Longitud del lugar donde se desactivó el taxímetro.                                                                                       |
| `dropoff_latitude`      | Latitud del lugar donde se desactivó el taxímetro.                                                                                        |
| `payment_type`          | Código numérico del método de pago (1=Credit card, 2=Cash, 3=No charge, 4=Dispute, 5=Unknown, 6=Voided trip).                             |
| `fare_amount`           | Tarifa calculada por tiempo y distancia según el taxímetro.                                                                               |
| `extra`                 | Cargos y recargos varios (p. ej., recargos por hora punta, nocturnos).                                                                    |
| `MTA_tax`               | Impuesto MTA de 0.50 aplicado automáticamente según la tarifa.                                                                            |
| `improvement_surcharge`  | Recargo de mejora de 0.30 aplicado al inicio del viaje (vigente desde 2015).                                                              |
| `tip_amount`            | Propina (poblado automáticamente para propinas con tarjeta; no incluye efectivo).                                                         |
| `tolls_amount`          | Total de peajes pagados en el viaje.                                                                                                      |
| `total_amount`          | Precio total cobrado al pasajero (no incluye propinas en efectivo).                                                                       |


# INSTRUCCIONES

En cada celda debes responder a la pregunta formulada, asegurándote de que el resultado queda guardado en la(s) variable(s) que por defecto vienen inicializadas a None. No se necesita usar variables intermedias, pero puedes hacerlo siempre que el resultado final del cálculo quede guardado exactamente en la variable que venía inicializada a None (debes reemplazar None por la secuencia de transformaciones necesarias, pero nunca cambiar el nombre de esa variable).

**No olvides borrar la línea `raise NotImplementedError()` de cada celda cuando hayas completado la solución de esa celda y quieras probarla.**

Después de cada celda evaluable verás una celda con código. Ejecútala (no modifiques su código) y te dirá si tu solución es correcta o no. Además de esas pruebas, se realizarán algunas más (ocultas) a la hora de puntuar el ejercicio, pero evaluar dicha celda es un indicador bastante fiable acerca de si realmente has implementado la solución correcta o no. Asegúrate de que, al menos, todas las celdas indican que el código es correcto antes de enviar el notebook terminado.

**No se debe redondear ninguna cantidad si no lo pide explícitamente el enunciado, ni aplicar ningún preprocesamiento que no pida el enunciado**

### Cada solución debe escribirse obligatoriamente en la celda habilitada para ello. Cualquier celda adicional que se haya creado durante el desarrollo deberá ser eliminada.

Si necesitas crear celdas auxiliares durante el desarrollo, puedes hacerlo pero debes asegurarte de borrarlas antes de entregar el notebook.

#### ENTORNO DE DATABRICKS
* Si estás usando Databricks Standard: debes editar tu cluster para usar DBR 18.1 que trae Spark 4.1.0
* Si estás usando Databricks Free Edition (donde el cluster es muy rápido): debes usar la versión 5 del entorno (Standard v5) con Spark 4.1.0. En ese caso, recuerda poner en cada lectura de Spark las opciones siguientes, poniendo el nombre de tu storage account:
  * `.option("adls.account.name", "nombre_de_tu_storage_account")`
  * `.option("fs.azure.account.key.nombre_de_tu_storage_account.dfs.core.windows.net", "tu_API_key_del_storage")`
* Recuerdra también que en Databricks Free Edition no se pueden cachear los DF (pero eso no implica que no debas dejar tu código preparado con los cache() en los sitios adecuados). Puedes inicialmente escribir tu solución y probar los tests, y antes de entregar, debes revisar los enunciados para añadir .cache() en las variables solicitadas, aunque no puedas probar el funcionamiento.

## Ejercicio 1 (1.5 puntos)

Con el fichero de taxis, se pide crear una función `read_and_clean_csv(path)` que realice las siguientes operaciones, y luego aplicar esa función para leer y limpiar los datos de los tres ficheros disponibles (enero, febrero y marzo de 2016), concatenando los resultados con `union` y guardando el resultado final en la variable `nyc_trips_df`, que debe quedar cacheado.

**NOTA IMPORTANTE**: si has instalado Spark en tu ordenador, debes configurar `spark.conf.set("spark.sql.session.timeZone", "UTC")` para que Spark lea las fechas con la zona horaria correcta, y no tener problemas al comprobar los resultados. Si estás usando Databricks, el Databricks Runtime 18.0 ya trae Spark 4.1.0 que tiene la zona horaria UTC por defecto, por lo que no tendrás que hacer esta configuración. Si estás usando Databricks Free Edition, debes usar versión del entorno 5 (Standard v5).

1. Leer los datos del fichero CSV, tratando de que Spark infiera el esquema.
2. Eliminar los viajes que tienen distancia 0 o negativa, precio total 0 o negativo, precio del taxímetro 0 o negativo,
o precio total menor estricto que el precio calculado por el taxímetro.
3. Eliminar aquellos viajes cuya distancia es estrictamente menor de 0.62 millas (1 km) o bien con duración estrictamente menor de 2 minutos, ya que es muy probable que se trate de viajes erróneos o mal registrados.
4. Como los viajes no poseen ningún identificador, añadir una columna llamada `trip_id` como `F.abs(F.xxhash64(F.concat("tpep_pickup_datetime", "tpep_dropoff_datetime")))`, que calcula un código hash (código numérico único a partir de cualquier objeto).
5. Añadir una columna llamada `trip_duration_s` con el resultado de calcular la duración del viaje en segundos, a partir de las columnas `tpep_pickup_datetime` y `tpep_dropoff_datetime`, usando la función `F.timestamp_diff(unit, start_col, end_col)` con el argumento `unit="SECOND"`.
5. Crear las columnas `"pickup_hour"` y `"dropoff_hour"` extrayendo la hora del día (solo la hora) de los campos respectivos.
6. Crear la columna `day_of_week` como el día de la semana de la hora de inicio del viaje (con `F.dayofweek(...)`, la cual devuelve un entero empezando en 1), traducida a una palabra en español. NOTA: la función devuelve un 1 para el domingo, un 2 para el lunes... hasta un 7 para el sábado, y por tanto, debe recategorizarse con `F.when` a `lunes`, `martes`, `miercoles`, ... `sabado`, `domingo` (no pongas tildes para evitar como norma general manejar en los DataFrames cadenas con tildes). PISTA: puedes hacerlo en dos pasos, con dos transformaciones withColumn encadenadas. En el primero creas la columna, y en el segundo, la recategorizas.
7. Crear la columna `speed` calculando la velocidad media del viaje en millas por hora redondeada a 2 cifras decimales (con `F.round(..., 2)`), a partir de la distancia y la duración calculada en segundos, teniendo en cuenta que 1 hora tiene 3600 segundos (recuerda que es posible dividir una columna entre un número).
8. Crear la columna `tip_perc`, como el porcentaje de propina (en tanto por uno), dividiendo la propina entre el precio total (`total_amount`), sin multiplicar por 100.

La función read_and_clean_csv debe resolverse en una única secuencia de transformaciones encadenadas.

```python
 !pip install folium==0.20.0
 dbutils.library.restartPython()
```

Output:
```
Requirement already satisfied: folium==0.20.0 in /local_disk0/.ephemeral_nfs/envs/pythonEnv-ed87b1b1-aaab-41a8-beb9-23616dccb43f/lib/python3.12/site-packages (0.20.0)
Requirement already satisfied: branca>=0.6.0 in /local_disk0/.ephemeral_nfs/envs/pythonEnv-ed87b1b1-aaab-41a8-beb9-23616dccb43f/lib/python3.12/site-packages (from folium==0.20.0) (0.8.2)
Requirement already satisfied: jinja2>=2.9 in /databricks/python3/lib/python3.12/site-packages (from folium==0.20.0) (3.1.6)
Requirement already satisfied: numpy in /databricks/python3/lib/python3.12/site-packages (from folium==0.20.0) (2.1.3)
Requirement already satisfied: requests in /databricks/python3/lib/python3.12/site-packages (from folium==0.20.0) (2.32.3)
Requirement already satisfied: xyzservices in /local_disk0/.ephemeral_nfs/envs/pythonEnv-ed87b1b1-aaab-41a8-beb9-23616dccb43f/lib/python3.12/site-packages (from folium==0.20.0) (2026.3.0)
Requirement already satisfied: MarkupSafe>=2.0 in /databricks/python3/lib/python3.12/site-p
```

<!-- grade_id: cell-e3d7010268c1a90b, points: 0, grade: false, solution: true -->
```python
from pyspark.sql import DataFrame as DF
from pyspark.sql import functions as F


taxi_csv_2016_01_path = "abfss://datos@matermcb001sta.dfs.core.windows.net/yellow_tripdata_2016-01.csv"

taxi_csv_2016_02_path = "abfss://datos@matermcb001sta.dfs.core.windows.net/yellow_tripdata_2016-02.csv"

taxi_csv_2016_03_path = "abfss://datos@matermcb001sta.dfs.core.windows.net/yellow_tripdata_2016-03.csv"

def read_and_clean_csv(path: str) -> DF:

    transformed_df = (
        spark.read
        .format("csv")
        .option("header", "true")
        .option("inferSchema", "true")
        .option("adls.account.name", "matermcb001sta")
        .option(
            "fs.azure.account.key.matermcb001sta.dfs.core.windows.net",
            "a+X6TcqNgcaSv+lTwr0jsVKuiFCRQ0UL5KfOcXV5b80TwDol3xW1VcLzvD7tkb+ZE5kyMxo8zxuz+AStaCTWOQ=="
        )

        .load(path)

        # duración del viaje en segundos
        .withColumn(
            "trip_duration_s",
            F.timestamp_diff(
                "SECOND",
                F.col("tpep_pickup_datetime"),
                F.col("tpep_dropoff_datetime")
            )
        )

        # filtros
        .filter(F.col("trip_distance") > 0)
        .filter(F.col("total_amount") > 0)
        .filter(F.col("fare_amount") > 0)
        .filter(F.col("total_amount") >= F.col("fare_amount"))
        .filter(F.col("trip_distance") >= 0.62)
        .filter(F.col("trip_duration_s") >= 120)

        # identificador único
        .withColumn(
            "trip_id",
            F.abs(
                F.xxhash64(
                    F.concat(
                        "tpep_pickup_datetime",
                        "tpep_dropoff_datetime"
                    )
                )
            )
        )

        # horas
        .withColumn(
            "pickup_hour",
            F.hour("tpep_pickup_datetime")
        )

        .withColumn(
            "dropoff_hour",
            F.hour("tpep_dropoff_datetime")
        )

        # día de la semana numérico
        .withColumn(
            "day_of_week",
            F.dayofweek("tpep_pickup_datetime")
        )

        # traducir día semana
        .withColumn(
            "day_of_week",
            F.when(F.col("day_of_week") == 1, "domingo")
            .when(F.col("day_of_week") == 2, "lunes")
            .when(F.col("day_of_week") == 3, "martes")
            .when(F.col("day_of_week") == 4, "miercoles")
            .when(F.col("day_of_week") == 5, "jueves")
            .when(F.col("day_of_week") == 6, "viernes")
            .when(F.col("day_of_week") == 7, "sabado")
        )

        # velocidad media
        .withColumn(
            "speed",
            F.round(
                (F.col("trip_distance") / F.col("trip_duration_s")) * 3600,
                2
            )
        )

        # porcentaje propina
        .withColumn(
            "tip_perc",
            F.col("tip_amount") / F.col("total_amount")
        )
    )

    return transformed_df


nyc_trips_df = (
    read_and_clean_csv(taxi_csv_2016_01_path)
    .union(read_and_clean_csv(taxi_csv_2016_02_path))
    .union(read_and_clean_csv(taxi_csv_2016_03_path))
.cache()
)


```

<!-- grade_id: cell-2631eae629064036, points: 1.5, grade: true, solution: false -->
```python
# Tiempo de ejecución esperado: entre 8 y 12 minutos
assert nyc_trips_df.is_cached
assert(nyc_trips_df.count() == 31088883)
assert( round(nyc_trips_df.select(F.mean("trip_distance").alias("trip_distance")).first().trip_distance, 2) == 5.18 )
assert( round(nyc_trips_df.select(F.mean("speed").alias("speed")).first().speed, 2) == 21.61 )
```

## Ejercicio 2 (1.5 puntos)

Vamos a leer un fichero con los eventos que han tenido lugar en Nueva York, con el objetivo de cruzarlos en ejercicios posteriores con los recorridos de los taxis.

Con el fichero de datos de eventos en Nueva York, se pide:

1. Leer los datos del fichero CSV en la variable, tratando de que Spark infiera el esquema, con separador ";".
2. Reemplazar las columnas `"Close Time"` y `"Create Time"` por el resultado de convertirlas a timestamp con `F.try_to_timestamp` con el formato `F.lit("MM/dd/yyyy hh:mm:ss a")`, que devuelve NULL cuando no se puede realizar la conversión porque el string de la celda no sigue el formato especificado.
    * Hay celdas que no siguen estrictamente el formato "MM/dd/yyyy hh:mm:ss a", sino que usan formato de hora en 24 horas sin AM/PM, como `"MM/dd/yyyy H:mm"` (por ejemplo, "11/04/2017 21:32" que no trae AM/PM y no especifica los segundos) Por tanto, en lugar de poner `F.try_to_timestamp` directamente, se debe usar `F.coalesce` poniendo en primer lugar la conversión con el formato con AM/PM, y si devuelve NULL, en segundo lugar usamos  `F.to_timestamp` (sin try) con el formato en 24 horas `"MM/dd/yyyy H:mm"`. Es decir, `F.coalesce(F.try_to_timestamp(...), F.to_timestamp(...))`. Esto debe hacerse tanto para `Close Time` como para `Create Time`
        * La función `F.coalesce(col1, col2, ...)` devuelve un objeto columna en el cual, en cada fila, estará el valor de la primera columna que no sea NULL que encuentre cuando va recorriendo en esa fila las columnas que le hemos pasado como argumentos (las irá recorriendo en el orden en el que se las hemos pasado).
3. Las columnas Longitude y Latitude no se han convertido correctamente porque hay filas mal formadas. Se pide reemplazar las columnas originales con el mismo nombre por su versión convertida a double, reemplazando primero las comas por puntos con `F.regexp_replace` y usando el objeto columna resultante para convertirlas a tipo double. Es decir, `F.regexp_replace(..., ",", ".").cast("double")`.

En ocasiones, la fecha del evento no coincide con la indicada en Create Time o Close Time, sino que se crea el evento con mucha antelación, y se indica la fecha exacta en la que tendrá lugar
en la columna `"Event Description"`. Por ejemplo: `Concert at The Beacon Theatre (Manhattan) Brandi Carlile, Thursday April 5th, 2018 thru Saturday April 7th, 2018`.
A veces incluso indica un rango de días en los que se repetirá el evento, pero en esos casos basta con extraer la primera fecha. Se pide:

4. Extraer el mes y el día del evento de la columna `"Event Description"` mediante la expresión regular indicada en la siguiente celda. Se debe usar con `F.regexp_extract`
para crear nuevas columnas `"event_month"`, `"event_day"`, y `"event_year"`, las tres de tipo string, pasando en cada caso a la función `regexp_extract` el argumento `idx` con valores 1, 2 y 3 respectivamente, para extraer el contenido del grupo de matching (paréntesis) primero, segundo y tercero de la expresión regular, que son el mes, el día del mes, y el año. Por último, crear una nueva columna `"event_date"` de tipo Date, concatenando las tres anteriores por ese orden, con un espacio en blanco, usando `F.try_to_date(F.concat_ws(" ", ...), "MMMM d yyyy")`. De esta forma está asegurado que la regex solamente captura la primera fecha, en caso de que haya un rango de fechas. Las columnas `event_month`, `event_day` y `event_year` deben eliminarse después de crear la columna `event_date`. Si no se encuentra la fecha en `Event Description`, la columna `event_date` habrá quedado a NULL, pero en ese caso debe reemplazarse por el valor de la fecha indicada en `Close Time` tras convertirlo a Date con `F.try_to_date`. Puedes usar `F.coalesce` para esto, poniendo en primer lugar la columna `event_date` y en segundo lugar la conversión a Date de `Close Time`.
    * **NOTA**: esta función se introdujo nueva en Spark 4.1.0 y por tanto, requiere que el Databricks Runtime seleccionado para el cluster sea el 18.0 que trae Spark 4.1.0.

La lectura y todas estas transformaciones (todas son withColumn) deben estar encadenadas, sin usar variables intermedias, y el resultado debe guardarse en la variable `raw_events_df`, sin cachear.
Para el último paso de crear event_date, puedes encadenar una transformación más a la que ya tenías, sin necesidad de crear una nueva variable intermedia, para reemplazar los posibles nulos de event_date por el resultado de convertir Close Time a Date.

<!-- grade_id: cell-28aaff8c4681ec34, points: 0, grade: false, solution: true -->
```python
from pyspark.sql import SparkSession, functions as F

regex_event_date = r'(?i)\b(January|February|March|April|May|June|July|August|September|October|November|December)\s+(\d{1}|\d{2})(?:st|nd|rd|th)?\s*,\s*(\d{4})\b'


csv_path = "abfss://datos@matermcb001sta.dfs.core.windows.net/511_NY_Events_raw.csv"

raw_events_df = (
    spark.read
    .format("csv")
    .option("header", "true")
    .option("inferSchema", "true")
    .option("sep", ";")
    .option("adls.account.name", "matermcb001sta")
    .option(
            "fs.azure.account.key.matermcb001sta.dfs.core.windows.net",
            "a+X6TcqNgcaSv+lTwr0jsVKuiFCRQ0UL5KfOcXV5b80TwDol3xW1VcLzvD7tkb+ZE5kyMxo8zxuz+AStaCTWOQ=="
        )
    .load(csv_path)

    # Close Time
    .withColumn(
        "Close Time",
        F.coalesce(
            F.try_to_timestamp(
                F.col("Close Time"),
                F.lit("MM/dd/yyyy hh:mm:ss a")
            ),
            F.to_timestamp(
                F.col("Close Time"),
                "MM/dd/yyyy H:mm"
            )
        )
    )

    # Create Time
    .withColumn(
        "Create Time",
        F.coalesce(
            F.try_to_timestamp(
                F.col("Create Time"),
                F.lit("MM/dd/yyyy hh:mm:ss a")
            ),
            F.to_timestamp(
                F.col("Create Time"),
                "MM/dd/yyyy H:mm"
            )
        )
    )

    # Latitude
    .withColumn(
        "Latitude",
        F.regexp_replace(
            F.col("Latitude"),
            ",",
            "."
        ).cast("double")
    )

    # Longitude
    .withColumn(
        "Longitude",
        F.regexp_replace(
            F.col("Longitude"),
            ",",
            "."
        ).cast("double")
    )

    # event_month
    .withColumn(
        "event_month",
        F.regexp_extract(
            F.col("Event Description"),
            regex_event_date,
            1
        )
    )

    # event_day
    .withColumn(
        "event_day",
        F.regexp_extract(
            F.col("Event Description"),
            regex_event_date,
            2
        )
    )

    # event_year
    .withColumn(
        "event_year",
        F.regexp_extract(
            F.col("Event Description"),
            regex_event_date,
            3
        )
    )

    # event_date
    .withColumn(
        "event_date",
        F.coalesce(
            F.try_to_date(
                F.concat_ws(
                    " ",
                    F.col("event_month"),
                    F.col("event_day"),
                    F.col("event_year")
                ),
                "MMMM d yyyy"
            ),
            F.try_to_date(F.col("Close Time"))
        )
    )

    # borrar auxiliares
    .drop(
        "event_month",
        "event_day",
        "event_year"
    )
)
```

<!-- grade_id: cell-1f6c21cebd5f6a91, points: 1.5, grade: true, solution: false -->
```python
assert(raw_events_df.count() == 13715)

# Casuística 1: event_date se ha extraído correctamente de Event Description
r1 = raw_events_df.where(F.col("Event Description").contains("Tuesday June 2nd"))\
                  .select(F.date_format("event_date", "yyyy-MM-dd").alias("event_date")).first()
assert(r1.event_date == "2020-06-02")

# Casuística 2: event_date no se ha podido extraer de Event Description, por lo que se ha deducido de Close Time
r2 = raw_events_df.where( (F.date_format("Close Time", "yyyy-MM-dd") == "2022-05-28") &
                          (F.col("Facility Name") == "Meadowlands Racetrack"))\
                  .select(F.date_format("event_date", "yyyy-MM-dd").alias("event_date")).first()
assert(r2.event_date == "2022-05-28")
```

## Ejercicio 3 (1.5 puntos)

1. En los casos donde la fecha se indica en `Event Description`, esa columna contiene también la hora de inicio y fin del evento (a veces sólo indica la hora de fin como "until 10:00 PM", pero no la de inicio). Suele venir en expresiones como `07:00 PM thru 10:00 PM` o `from 08:00 PM until 11:00 PM` (en ocasiones A.M. ó P.M. en lugar de AM ó PM). Se pide crear nuevas columnas `"Extracted Start Time"` y `"Extracted End Time"` **de tipo string**, que contengan la hora de inicio y fin del evento en formato "11:00 PM" o "10:30 P.M.", extrayéndola de la columna `Event Description` mediante las expresiones regulares indicadas en la siguiente celda. La primera regex funciona cuando están presentes ambas horas, y la segunda cuando solo está presente la hora de fin (precedida de la palabra "until"). En este último caso, la columna `"Extracted Start Time"` quedará en blanco.
  * PISTA: para crear la columna `"Extracted Start Time"` la primera vez, utiliza `F.regexp_extract` con `regex_event_time_both` e índice 1, y para `"Extracted End Time"` utilizar `F.regexp_extract` con `regex_event_time_both` e índice 2 para intentar extraer la hora de fin, en un primer intento.
  * Como habrá casos donde no se habrá podido extraer adecuadamente la hora de fin por no contener también la hora de inicio, reemplazar `"Extracted End Time"`que acabamos de crear y tenía valor `""` en esos casos (utilizando `F.when` para comprobar si la columna `"Extracted End Time"` es igual a la cadena vacía `""`), por el resultado de aplicar `F.regexp_extract` con `regex_event_time_until` e índice 1. En caso contrario (si pudimos extraer adecuadamente `"Extracted End Time"`), se debe mantener el valor original de `"Extracted End Time"`.
  * Por último, haremos un último ajuste a `Extracted Start time`. Reemplaza (de nuevo) dicha columna por el resultado de eliminar el "." que existe en algunos casos en ambas columnas, debido a que la hora puede venir como "10:30 P.M." en lugar de "10:30 PM", utilizando `F.regexp_replace` para reemplazar el carácter "." por la cadena vacía `""`.

2. Crear una nueva columna `"event_end_ts"` que pase a timestamp el resultado de concatenar la columna `"event_date"` (obtenida en el paso anterior) con la columna `"Extracted End Time"`, separadas por un espacio en blanco, con `F.try_to_timestamp(F.concat_ws(...), "yyyy-MM-dd h:mm a")`. Repetir la operación para crear una nueva columna `"event_start_ts"`, concatenando `"event_date"` con `"Extracted Start Time"`. Los posibles nulos generados los trataremos después.
    * CUIDADO: una vez creadas ambas columnas, si la hora de finalización es anterior a la de inicio, habrá que sumar un día a la fecha de finalización para que tenga sentido (el evento termina al día siguiente). Para ello, después de haber creado `event_end_ts`, encadenar otra transformación para reemplazarla por el resultado de aplicar `F.when` para comprobar si `"event_end_ts"` es menor que `"event_start_ts"`, y en ese caso reemplazar `"event_end_ts"` por el resultado de añadir un día a `"event_end_ts"` con `F.date_add("event_end_ts", 1)`. Si no se cumple esa condición, mantener el valor original de `"event_end_ts"`.
3. Tras esto, en los casos en los que `event_end_ts` fuese nulo (no se pudo extraer la hora de fin con la regex, la cual devolvió cadena vacía, y por eso su conversión a timestamp resulta en nulo), establecer `event_end_ts` como el valor de `Close Time` (convertido a timestamp en un apartado previo). Esto puede hacerse usando `F.coalesce` (consulta la documentación de la función). En los casos en los `event_start_ts` fuese nulo (no se pudo extraer la hora de inicio con la regex, devolvió cadena vacía y por eso su conversión a timestamp resulta en nulo), debemos crear `event_start_ts` como el resultado de restar 2 horas a `"event_end_ts"`, incluso para los casos donde `event_end_ts` fue copiado de `Close TIme`, utilizando para ello `F.timestamp_add("HOUR", -2, "event_start_ts")`. Estamos asumiendo que los eventos duran dos horas para deducir la hora de inicio. Aquí puede usarse también `F.coalesce`.

4. Eliminar aquellas filas en las que `event_date`, `event_start_ts` o `event_end_ts` sean nulos cualquiera de los tres.
5. Eliminar filas duplicadas considerando únicamente las columnas `Event Description` y `event_date` simultáneamente.
6. El resultado de todas estas transformaciones (encadenadas) debe guardarse en un nuevo DataFrame llamado `events_df`, que debe quedar cacheado. No debe utilizarse ninguna variable intermedia, ni asignaciones intermedias del estilo `event_df = event_df.funcion(...)` (serán penalizadas).


<!-- grade_id: cell-1d391f86d1e52f19, points: 0, grade: false, solution: true -->
```python
regex_event_time_both = r'(?i)(?:from\s+)?(\d{1,2}:\d{2}\s*(?:AM|PM|A\.M\.|P\.M\.))\s*(?:thru|until|to|-)\s*(\d{1,2}:\d{2}\s*(?:AM|PM|A\.M\.|P\.M\.))'
regex_event_time_until = r'(?i)\s*until\s*(\d{1,2}:\d{2}\s*(?:AM|PM|A\.M\.|P\.M\.))'

events_df = (
    raw_events_df

    # Extracted Start Time
    .withColumn(
        "Extracted Start Time",
        F.regexp_extract(
            F.col("Event Description"),
            regex_event_time_both,
            1
        )
    )

    # Extracted End Time
    .withColumn(
        "Extracted End Time",
        F.regexp_extract(
            F.col("Event Description"),
            regex_event_time_both,
            2
        )
    )

    # casos "until"
    .withColumn(
        "Extracted End Time",
        F.when(
            F.col("Extracted End Time") == "",
            F.regexp_extract(
                F.col("Event Description"),
                regex_event_time_until,
                1
            )
        ).otherwise(
            F.col("Extracted End Time")
        )
    )

    # quitar puntos
    .withColumn(
        "Extracted Start Time",
        F.regexp_replace(
            F.col("Extracted Start Time"),
            r"\.",
            ""
        )
    )

    .withColumn(
        "Extracted End Time",
        F.regexp_replace(
            F.col("Extracted End Time"),
            r"\.",
            ""
        )
    )

    # event_end_ts
    .withColumn(
        "event_end_ts",
        F.try_to_timestamp(
            F.concat_ws(
                " ",
                F.col("event_date"),
                F.col("Extracted End Time")
            ),
            F.lit("yyyy-MM-dd h:mm a")
        )
    )

    # event_start_ts
    .withColumn(
        "event_start_ts",
        F.try_to_timestamp(
            F.concat_ws(
                " ",
                F.col("event_date"),
                F.col("Extracted Start Time")
            ),
            F.lit("yyyy-MM-dd h:mm a")
        )
    )

    # si termina al día siguiente
    .withColumn(
        "event_end_ts",
        F.when(
            F.col("event_end_ts") < F.col("event_start_ts"),
            F.date_add(
                F.col("event_end_ts"),
                1
            )
        ).otherwise(
            F.col("event_end_ts")
        )
    )

    # rellenar event_end_ts
    .withColumn(
        "event_end_ts",
        F.coalesce(
            F.col("event_end_ts"),
            F.col("Close Time")
        )
    )

    # rellenar event_start_ts
    .withColumn(
        "event_start_ts",
        F.coalesce(
            F.col("event_start_ts"),
            F.expr("event_end_ts - INTERVAL 2 HOURS"
        )
        )
    )

    # eliminar nulos
    .filter(
        F.col("event_date").isNotNull()
    )

    .filter(
        F.col("event_start_ts").isNotNull()
    )

    .filter(
        F.col("event_end_ts").isNotNull()
    )

    # eliminar duplicados
    .dropDuplicates(
        ["Event Description", "event_date"]
    )

    .cache()
)
```

<!-- grade_id: cell-4d4daef1fb1c38e7, points: 1.5, grade: true, solution: false -->
```python
assert(events_df.count() == 12282)

# Casuística 1: event start y event end tienen valor extraído de Event Description
r1 = events_df.where(F.col("Event Description").contains("Tuesday June 2nd"))\
              .select(F.date_format("event_start_ts", "yyyy-MM-dd HH:mm:ss").alias("event_start_ts"),
                 F.date_format("event_end_ts", "yyyy-MM-dd HH:mm:ss").alias("event_end_ts")).first()
assert(r1.event_start_ts == "2020-06-02 19:00:00" and r1.event_end_ts == "2020-06-02 23:00:00")

# Casuística 2: event start no se pudo extraer de Event Description, se ha establecido restando 2 horas a event end, el cual sí se ha extraído de Event Description
r2 = events_df.where("event_date = '2020-03-07' and `Facility Name`='Meadowlands Racetrack'")\
              .select(F.date_format("event_start_ts", "yyyy-MM-dd HH:mm:ss").alias("event_start_ts"),
                      F.date_format("event_end_ts", "yyyy-MM-dd HH:mm:ss").alias("event_end_ts")).first()
assert(r2.event_start_ts == "2020-03-07 21:00:00" and r2.event_end_ts == "2020-03-07 23:00:00")

# Casuística 3: event end no se pudo extraer de Event Description, se ha establecido con el valor de Close Time, y event start se ha establecido restando 2 horas a event end
r3 = events_df.where("`Event Description` like '%The Revivalists%' and event_date = '2020-03-21'")\
              .select(F.date_format("event_start_ts", "yyyy-MM-dd HH:mm:ss").alias("event_start_ts"),
                      F.date_format("event_end_ts", "yyyy-MM-dd HH:mm:ss").alias("event_end_ts")).first()
assert(r3.event_start_ts == "2020-03-21 21:17:19" and r3.event_end_ts == "2020-03-21 23:17:19")
```

## Ejercicio 4 (1 punto)

Vamos a visualizar en un mapa los puntos de recogida de dos muestras distintas de viajes. Para ello, primero debemos subir al Workspace el script auxiliar `aux_func.py` que contiene varias funciones de visualización. Debes subirlo a la misma carpeta donde está este notebook, para que la importación del módulo aux_func en una celda funcione directamente.

Una vez hecho esto, partiendo de `nyc_trips_df`, se pide:

* Viajes que se han detectado como outliers por estar el punto de recogida fuera del área geográfica de la ciudad de Nueva York. Debes primero calcular quiénes son estos viajes, y guardarlos en la variable `outlier_locations_df`. El área de la ciudad de Nueva York se puede acotar como un rectángulo cuyas coordenadas tienes en la celda. Los outliers tendrán longitud *estrictamente menor* que la longitud mínima o *estrictamente mayor* que la longitud máxima, y lo mismo con la latitud.
* Viajes que considerados inliers (dentro del área geográfica de Nueva York). Debes primero calcular quiénes son estos viajes, y guardarlos en la variable `inlier_locations_df`, la cual debe quedar cacheada.
* Una vez que hayas calculado ambos DF, utiliza la transformación `sample` (sin reemplazamiento y con semilla 12345) para quedarte con una muestra del 1% de los outliers, y con una muestra del 0.001% de los inliers, y convertir ambas muestras a Pandas DataFrames con `toPandas()`, guardando el resultado en las variables `outliers_nyc_pd` e `inliers_nyc_pd` respectivamente. Recuerda que el argumento de sample es el porcentaje medido en tanto por uno (número entre 0.0 y 1.0).
* El motivo de quedarnos sólo con una muestra tan pequeña de los inliers es que hay muchísimos viajes dentro del área geográfica de Nueva York, y si intentamos visualizarlos todos en el mapa, tardará mucho tiempo en renderizarse.
* Para visualizarlos, simplemente invoca a la función `visualiza_viajes` que tienes ya hecha, sin modificar su código, pasando como argumentos los dos DF de Pandas. Asegúrate de instalar el paquete `folium` haciendo `!pip install folium==0.20.0` en el propio notebook. El objeto devuelto por la función debe almacenarse en la variable `nyc_map`, y se mostrará automáticamente al ser la última línea de la celda.


<!-- grade_id: cell-a4a0626314f2bce0, points: 0, grade: false, solution: true -->
```python
# Tiempo de ejecución esperado: 4 minutos

from aux_func import visualiza_viajes
!pip install folium==0.20.0

nyc_boundaries = {
    "min_longitude": -74.15,
    "max_longitude": -73.7004,
    "min_latitude": 40.5774,
    "max_latitude": 40.9176
}

outlier_locations_df = nyc_trips_df.filter(
    (F.col("pickup_longitude") < nyc_boundaries["min_longitude"]) |
    (F.col("pickup_longitude") > nyc_boundaries["max_longitude"]) |
    (F.col("pickup_latitude") < nyc_boundaries["min_latitude"]) |
    (F.col("pickup_latitude") > nyc_boundaries["max_latitude"])
.cache()
)

inlier_locations_df = (nyc_trips_df.filter(
        (F.col("pickup_longitude") >= nyc_boundaries["min_longitude"]) &
        (F.col("pickup_longitude") <= nyc_boundaries["max_longitude"]) &
        (F.col("pickup_latitude") >= nyc_boundaries["min_latitude"]) &
        (F.col("pickup_latitude") <= nyc_boundaries["max_latitude"])
)
.cache()
)

outliers_nyc_pd = (
    outlier_locations_df
    .sample(False, 0.01, 12345)
    .toPandas()
)

inliers_nyc_pd = (
    inlier_locations_df
    .sample(False, 0.00001, 12345)
    .toPandas()
)

nyc_map = visualiza_viajes(outliers_nyc_pd, inliers_nyc_pd, nyc_boundaries)
```

Output:
```
Requirement already satisfied: folium==0.20.0 in /local_disk0/.ephemeral_nfs/envs/pythonEnv-ed87b1b1-aaab-41a8-beb9-23616dccb43f/lib/python3.12/site-packages (0.20.0)
Requirement already satisfied: branca>=0.6.0 in /local_disk0/.ephemeral_nfs/envs/pythonEnv-ed87b1b1-aaab-41a8-beb9-23616dccb43f/lib/python3.12/site-packages (from folium==0.20.0) (0.8.2)
Requirement already satisfied: jinja2>=2.9 in /databricks/python3/lib/python3.12/site-packages (from folium==0.20.0) (3.1.6)
Requirement already satisfied: numpy in /databricks/python3/lib/python3.12/site-packages (from folium==0.20.0) (2.1.3)
Requirement already satisfied: requests in /databricks/python3/lib/python3.12/site-packages (from folium==0.20.0) (2.32.3)
Requirement already satisfied: xyzservices in /local_disk0/.ephemeral_nfs/envs/pythonEnv-ed87b1b1-aaab-41a8-beb9-23616dccb43f/lib/python3.12/site-packages (from folium==0.20.0) (2026.3.0)
Requirement already satisfied: MarkupSafe>=2.0 in /databricks/python3/lib/python3.12/site-p
```

<!-- grade_id: cell-73df5ce9f4156709, points: 1, grade: true, solution: false -->
```python
assert(outlier_locations_df.count() == 440553)
assert(inlier_locations_df.count() == 30648330)
```

## Ejercicio 5 (3 puntos)

Vamos a comprobar si los eventos causan mayor afluencia de taxis en sus inmediaciones durante el horario del evento, y estudiaremos de dónde proviene mayoritariamente la demanda.

Partiendo de `events_df` y de `inlier_locations_df`, se pide:
* Filtrar `events_df` para quedarnos solamente con los eventos cuya fecha está entre el 1 de enero y el 31 de marzo de 2016, ambos inclusive, y que además tengan lugar a partir de las 18 horas inclusive (en la hora de inicio del evento). El resultado se guardará en `events_df_filtered`.
* Filtrar para quedarnos solamente con los viajes en los que la hora de llegada sea a partir de las 17 horas inclusive (dejamos un marge con respecto a la hora de inicio, que puede ser a partir de las 18, para incluir viajes que hayan finalizado poco antes del evento). El resultado se guardará en `nyc_trips_filtered`, el cual debe quedar cacheado.
* Filtrar `nyc_trips_filtered` para quedarnos solamente con los viajes que tengan el punto de llegada a una distancia de Manhattan menor o igual a 100 metros de la ubicación del evento. Para esto, primero debes crear una función `manhattan_distance_meters` que calcule la distancia de Manhattan en metros entre dos puntos geográficos dados por sus latitudes y longitudes, y devuelva un objeto Column con esa distancia.
  * Esta función debe recibir solamente 4 argumentos, que serán 4 strings con los nombres de las columnas latitud, longitud de un punto y latitud, longitud del otro punto, y debe operar todo el rato con objetos Column que no están ligados a ningún DataFrame (por eso no recibe ningún DF como argumento).
  * Para el cálculo de la distancia de Manhattan, puedes calcular por separado la componente Norte-Sur (diferencia de latitud) y la componente Este-Oeste como sigue, y devolver el objeto Column con la suma de ambas.
    * Diferencia norte-sur en metros: valor absoluto de la diferencia de latitudes (tras expresar primero ambas latitudes en radianes con `F.radians(colname)`), multiplicada por el radio de la Tierra en metros (6371000 m). Es decir, abs(lat2_rad - lat1_rad) * 6371000. El sufijo "_rad" indica que las magnitudes han sido previamente convertidas a radianes con `F.radians(colname)`.
    * Diferencia este-oeste en metros: valor absoluto de la diferencia de longitudes (tras expresar primero ambas longitudes en radianes con `F.radians(colname)`), multiplicada por el **coseno de la media de las dos latitudes (en radianes)** de ambos puntos y multiplicado por el radio de la Tierra en metros. Es decir, abs(lon2_rad - lon1_rad) * cos((lat1_rad + lat2_rad) / 2) * 6371000.
    * Recuerda que los objetos Column en Spark pueden operarse con sumas, restas, multiplicaciones, y que existe `F.abs`, `F.radians` y `F.cos` que reciben columnas como argumento y devuelven columnas.
* Una vez creada la función, filtrar `nyc_trips_filtered` para quedarnos solamente con los viajes que tengan el punto de llegada a una distancia de Manhattan menor o igual a 100 metros de la ubicación del evento, haciendo un join entre `nyc_trips_filtered` y el resultado de quitar duplicados por longitud y latitud a `events_df_filtered`, usando como condición de join una llamada a la distancia de Manhattan entre el punto de llegada del viaje y la ubicación del evento, que debe ser menor o igual a 100 metros.
  * OJO: estamos usando join para "filtrar" porque no nos interesan las columnas de la derecha... ¿qué tipo de join necesitas para no añadir columnas? El resultado se guardará en `nyc_trips_cercanos_sedes`, el cual debe quedar cacheado. Estos viajes son solamente los que terminan cerca de las ubicaciones de algún evento, pero sin tener en cuenta la hora del evento, por lo que todavía no sabemos si realmente han ido a esos eventos o simplemente han coincidido en la ubicación pero no en el horario. Será útil para comparar posteriormente con los viajes que sí coinciden también en el horario del evento, para ver si realmente hay un aumento de la demanda en las inmediaciones de los eventos durante su horario.
* Finalmente, partiendo de `nyc_trips_cercanos_sedes`, quedarnos solamente con los viajes que llegan a las inmediaciones de un evento (los del paso anterior) y que además llegan durante el horario del evento, es decir, que la diferencia en minutos entre la hora de llegada del viaje y la hora de inicio del evento esté entre 0 y 20 minutos. Para esto, se debe hacer un join entre `nyc_trips_cercanos_sedes` y `events_df_filtered`, usando como condición de join que la diferencia en minutos entre la hora de llegada del viaje (`tpep_dropoff_datetime`) y la hora de inicio del evento (`event_start_ts`) esté entre 0 y 20 minutos. No queremos perder ningún viaje, solo saber si coincide con el horario de algún evento o no, por lo que el tipo de join debe ser el adecuado para esto.
  * PISTA: la condición de join es que `F.timestamp_diff("MINUTE", ..., ...).between(..., ...)`. En el DF resultante deben quitarse los duplicados por `trip_id`, y después deben rellenarse los valores nulos en la columna "Event type" por el valor "No event" (casos donde no se encontró matching al ser un viaje que terminó cerca de la sede pero no en la franja horaria adecuada). El DF resultante se guardará en `nyc_trips_eventos`, el cual debe quedar cacheado. En este DataFrame ya sí podemos asumir que los viajes que aparecen allí han ido a esos eventos, porque han coincidido tanto en la ubicación como en el horario.

<!-- grade_id: cell-28f171b0d1829431, points: 0, grade: false, solution: true -->
```python
def manhattan_distance_meters(lat1_col: str, lon1_col: str, lat2_col: str, lon2_col: str):

    # coordenadas en radianes
    lat1_rad = F.radians(F.col(lat1_col))
    lon1_rad = F.radians(F.col(lon1_col))

    lat2_rad = F.radians(F.col(lat2_col))
    lon2_rad = F.radians(F.col(lon2_col))

    # componente norte-sur
    north_south = (
        F.abs(lat2_rad - lat1_rad) * 6371000
    )

    # componente este-oeste
    east_west = (
        F.abs(lon2_rad - lon1_rad) *
        F.cos((lat1_rad + lat2_rad) / 2) *
        6371000
    )

    # distancia Manhattan
    return north_south + east_west


events_df_filtered = (
    events_df.filter(
        (F.col("event_date") >= "2016-01-01") &
        (F.col("event_date") <= "2016-03-31") &
        (F.hour("event_start_ts") >= 18)
    )
    .cache()
)

nyc_trips_filtered = (
    inlier_locations_df.filter(
        F.hour("tpep_dropoff_datetime") >= 17
    )
    .cache()
)

   

nyc_trips_cercanos_sedes = (
    nyc_trips_filtered.join(
        events_df_filtered.dropDuplicates(
            ["Latitude", "Longitude"]
        ),
        manhattan_distance_meters(
            "dropoff_latitude",
            "dropoff_longitude",
            "Latitude",
            "Longitude"
        ) <= 100,
        "left_semi"
    )
    .cache()
)

nyc_trips_eventos = (
    nyc_trips_cercanos_sedes
    .join(
        events_df_filtered,

        F.timestamp_diff(
            "MINUTE",
            F.col("tpep_dropoff_datetime"),
            F.col("event_start_ts")
        ).between(0, 20),

        "left"
    )
    .orderBy("Event type")

    .dropDuplicates(["trip_id"])

    .fillna({
        "Event type": "No event"
    })

    .cache()
)

```

<!-- grade_id: cell-07b642411f9ffd4d, points: 3, grade: true, solution: false -->
```python
# Tiempo de cálculo esperado: alrededor de 5 minutos

# Primero probamos la función distancia de Manhattan
data = [
    (40.0, -74.0, 40.0, -73.999),
    (40.0, -74.0, 40.001, -74.0),
]
columns = ["pickup_latitude", "pickup_longitude", "dropoff_latitude", "dropoff_longitude"]
df_test = spark.createDataFrame(data, columns)

# Objeto columna con la distancia de Manhattan
manhattan_distance_col = manhattan_distance_meters(
    "pickup_latitude", "pickup_longitude",
    "dropoff_latitude", "dropoff_longitude")

rows = df_test.withColumn("manhattan", manhattan_distance_col).sort("manhattan").collect()
assert(round(rows[0].manhattan, 1) == 85.2)
assert(round(rows[1].manhattan, 1) == 111.2)

# Probamos los joins
assert(len(nyc_trips_cercanos_sedes.columns) == len(nyc_trips_filtered.columns))

assert(nyc_trips_filtered.count() == 12354050)
assert(events_df_filtered.count() == 169)
assert(nyc_trips_cercanos_sedes.count() == 184492)
assert(nyc_trips_eventos.where(F.col("Event type") != "No event").count() == 18057)

```

## Ejercicio 6 (1 punto)

Partiendo de `nyc_trips_eventos`, se pide calcular:
* Calcular estadísticas de los viajes agrupados por hora de llegada (`dropoff_hour`) y por tipo de evento (`Event type`). El DF resultante debe tener tantas filas como horas de llegada (y puesto que el DF de partida solo tiene horas de llegada a partir de las 17, tendría que tener 7 filas), y tantas columnas como tipos de eventos, multiplicadas por el número de agregaciones que vamos a realizar, más una (que es la hora de llegada). Las agregaciones que vamos a calcular para cada grupo son:
  * Número de viajes en cada hora y por cada tipo de evento en una columna llamada `num_trips`.
  * Distancia media de los viajes en cada hora y por cada tipo de evento, redondeada a dos decimales, en una columna llamada `avg_distance`.
  * Duración media en segundos de los viajes en cada hora y por cada tipo de evento, redondeada a dos decimales, en una columna llamada `avg_duration_s`.
  * Porcentaje de viajes pagados con tarjeta (payment_type = 1) en cada hora y por cada tipo de evento. Se puede hacer el promedio del resultado de comparar payment_type con el valor 1, convirtiendo el resultado a "int" con la función `cast("int")` para que se le pueda aplicar la media. El resultado de la media debe redondearse a dos decimales, en una columna llamada `card_payment_perc`.
  * Media del porcentaje de propina del viaje (tip_perc) en cada hora y por cada tipo de evento, considerando solamente los viajes pagados con tarjeta (payment_type = 1). Para esto, puedes usar `F.when` para quedarte solamente con el valor de `tip_perc` en los casos donde `payment_type` es igual a 1, y en caso contrario (`otherwise`) quedarte con `None` para que no se tengan en cuenta esos casos en el cálculo de la media. El resultado de esta agregación se guardará en una columna llamada `avg_tip_perc`. Es decir, `F.round(F.mean(F.when(...., F.col("tip_perc")).otherwise(None)), 2).alias(...)`.
* Tras las agregaciones, ordenar el DF obtenido ascendentemente por la hora de llegada.

El DF resultante debe quedar guardado en la variable `trips_statistics_df`.

Aunque no sea evaluable, analiza las siguientes cuestiones:
* ¿Por qué hay tantas casillas vacías (con valor nulo) en el resultado?
* ¿Qué tipo de eventos tienen más viajes asociados, y cuáles menos?
* ¿Hay alguna hora del día donde haya más viajes asociados a eventos que en otras horas?

<!-- grade_id: cell-4bb169603b77a19a, points: 0, grade: false, solution: true -->
```python
trips_statistics_df = (
    nyc_trips_eventos

    .groupBy("dropoff_hour")
    .pivot("Event type")
    .agg(
        # número viajes
        F.count("trip_id").alias("num_trips"),
        # distancia media
        F.round(
            F.avg("trip_distance"),
            2
        ).alias("avg_distance"),
        # duración media
        F.round(
            F.avg("trip_duration_s"),
            2
        ).alias("avg_duration_s"),
        # porcentaje pago tarjeta
        F.round(
            F.avg(
                (
                    F.col("payment_type") == 1
                ).cast("int")
            ),
            2
        ).alias("card_payment_perc"),
        # media tip_perc solo tarjeta
        F.round(
            F.mean(
                F.when(
                    F.col("payment_type") == 1,
                    F.col("tip_perc")
                ).otherwise(None)
            ),
            2
        ).alias("avg_tip_perc")
    )
    .orderBy("dropoff_hour")

    .cache()
)
```

<!-- grade_id: cell-62df1a60a3376f39, points: 1, grade: true, solution: false -->
```python
event_types_cols = ["No event", "basketball game", "concert", "hockey game", "wrestling match", "wrestling",
                    "boxing", "christmas spectacular", "other special event", "tennis tournament", "special event", "show"]
aggregations = ["num_trips", "avg_distance", "avg_duration_s", "card_payment_perc", "avg_tip_perc"]
combinations = [f"{event}_{agg}" for event in event_types_cols for agg in aggregations]
assert(set(trips_statistics_df.columns) == set(["dropoff_hour"] + combinations)) # comprobamos nombres de columnas

trips_list = trips_statistics_df.collect()
assert(trips_list[0].dropoff_hour == 17 and trips_list[0]["No event_num_trips"] == 28245)
assert(trips_list[1].dropoff_hour == 18 and trips_list[1]["basketball game_num_trips"] == 118)
```

## Ejercicio 7 (1 punto)

Partiendo de nuevo de `nyc_trips_eventos`, calcular las siguientes estadísticas agregadas, agrupando solamente por tipo de evento (`Event type`), sin tener en cuenta la hora de llegada. El resultado se guardará en `trips_agg_statistics_df`, y debe quedar cacheado:
* Número de viajes por cada tipo de evento en una columna llamada `num_trips`.
* Distancia media de los viajes por cada tipo de evento, redondeada a dos decimales, en una columna llamada `avg_distance`.
* Distancia mediana de los viajes por cada tipo de evento, redondeada a dos decimales, en una columna llamada `med_distance` (recuerda `F.median`).
* Duración media en segundos de los viajes por cada tipo de evento, redondeada a dos decimales, en una columna llamada `avg_duration_s`.
* Duración mediana en segundos de los viajes por cada tipo de evento, redondeada a dos decimales, en una columna llamada `med_duration_s`.
* Porcentaje de viajes pagados con tarjeta (payment_type = 1) por cada tipo de evento, redondeado a dos decimales, en una columna llamada `card_payment_perc`. Aplica las mismas reglas que en el ejercicio anterior para calcular este porcentaje, convirtiendo la columna booleana resultante de la comparación en una columna de tipo int (cuyos valores serán 1 y 0) para poder aplicar la media.
* Promedio del porcentaje de propina sobre el total del viaje (tip_perc) por cada tipo de evento, considerando solamente los viajes pagados con tarjeta (payment_type = 1), redondeado a dos decimales, en una columna llamada `avg_tip_perc`. Aplica las mismas reglas que en el ejercicio anterior para calcular este porcentaje.

Analiza visualmente la diferencia entre cada media y su mediana. A la vista de las medianas, ¿hay tanta variación entre grupos como indican la medias?

<!-- grade_id: cell-8f0ca5f98482284f, points: 0, grade: false, solution: true -->
```python
trips_agg_statistics_df = (
    nyc_trips_eventos
    .groupBy("Event type")
    .agg(
        # número viajes
        F.count("trip_id").alias("num_trips"),
        # distancia media
        F.round(
            F.avg("trip_distance"),
            2
        ).alias("avg_distance"),
        # distancia mediana
        F.round(
            F.median("trip_distance"),
            2
        ).alias("med_distance"),
        # duración media
        F.round(
            F.avg("trip_duration_s"),
            2
        ).alias("avg_duration_s"),
        # duración mediana
        F.round(
            F.median("trip_duration_s"),
            2
        ).alias("med_duration_s"),
        # porcentaje pago tarjeta
        F.round(
            F.avg(
                (
                    F.col("payment_type") == 1
                ).cast("int")
            ),
            2
        ).alias("card_payment_perc"),
        # promedio propina solo tarjeta
        F.round(
            F.mean(
                F.when(
                    F.col("payment_type") == 1,
                    F.col("tip_perc")
                ).otherwise(None)
            ),
            2
        ).alias("avg_tip_perc")
    )
    .cache()
)
```

<!-- grade_id: cell-1cbbbdca1f6c6e5f, points: 1, grade: true, solution: false -->
```python
trips_agg_statistics_pd = trips_agg_statistics_df.toPandas()
assert(trips_agg_statistics_pd.loc[trips_agg_statistics_pd["Event type"] == "No event", "num_trips"].values[0] == 166428)
assert(trips_agg_statistics_pd.loc[trips_agg_statistics_pd["Event type"] == "concert", "avg_distance"].values[0] in [1.87, 1.88])
assert(trips_agg_statistics_pd.loc[trips_agg_statistics_pd["Event type"] == "boxing", "card_payment_perc"].values[0] == 0.38)
```

Output:
```
[0;31m---------------------------------------------------------------------------[0m
[0;31mAssertionError[0m                            Traceback (most recent call last)
File [0;32m<command-6732434685386972>, line 3[0m
[1;32m      1[0m trips_agg_statistics_pd [38;5;241m=[39m trips_agg_statistics_df[38;5;241m.[39mtoPandas()
[1;32m      2[0m [38;5;28;01massert[39;00m(trips_agg_statistics_pd[38;5;241m.[39mloc[trips_agg_statistics_pd[[38;5;124m"[39m[38;5;124mEvent type[39m[38;5;124m"[39m] [38;5;241m==[39m [38;5;124m"[39m[38;5;124mNo event[39m[38;5;124m"[39m, [38;5;124m"[39m[38;5;124mnum_trips[39m[38;5;124m"[39m][38;5;241m.[39mvalues[[38;5;241m0[39m] [38;5;241m==[39m [38;5;241m166428[39m)
[0;32m----> 3[0m [38;5;28;01massert[39;00m(trips_agg_statistics_pd[38;5;241m.[39mloc[trips_agg_statistics_pd[[38;5;124m"[39m[38;5;124mEvent type[39m[38;5;124m"[39m] [38;5;241m==[39m [38;5;124m"[39m[38;5;124mconcert[39m[38;5;124m"[39m, [38;5
```